```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment3.ipynb")
```

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Probability, Statistics and Distances Homework

Please complete the following assignment. Run all cells and submit your completed notebook through Gradescope.

This assignment covers the following topics:

- Frequentist vs Bayesian probability
- Definition of probability
- Conditional probability and Bayes theorem
- Random variables
- Expectation, variance, and covariance
- Binomial distribution
- Poisson distribution
- Gaussian distribution
- Jaccard Index
- Euclidean Distances

<!-- BEGIN QUESTION -->

#### 1. Which of the following statements best describes the Bayesian interpretation of probability?

    A) Probability is the long-run relative frequency of an event occurring in repeated trials.
    B) Probability represents a degree of belief or confidence in the occurrence of an event.
    C) Probability is the ratio of favorable outcomes to the total number of equally likely outcomes.
    D) Probability is the measure of the likelihood that a random variable takes a specific value.

    Write a function `ans1` and return your answer as a string, e.g., "A".

```python
def ans1():
    return 'A'
print(ans1())
```

Output:
```
A
```

<!-- END QUESTION -->

#### 2. Write a function `roll_die(n:int) -> list` that simulates rolling a six-sided die `n` times and returns a `list` of outcomes.

```python
import numpy as np
import random
def roll_die(n: int) -> list:
    outcomes = []
    for roll in range(n):
        outcome = random.randint(1, 6)
        outcomes.append(outcome)
    return outcomes

print(roll_die(6))

    
```

Output:
```
[2, 6, 1, 6, 4, 6]
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

#### 3. Suppose that 1% of a population has a particular disease. A test for the disease correctly identifies  diseased individuals 99% of the time (sensitivity), and correctly identifies non-diseased individuals  95% of the time (specificity). If a randomly selected individual tests positive, what is the probability that they actually have the disease?
    Return your answer as a float between 0 and 1.

```python
def answer3():
    P_D = 0.01
    P_not_D = 0.99
    P_T_pos_given_D = 0.99
    P_T_pos_given_not_D = 1 - 0.95
    P_D_given_T_pos = (P_T_pos_given_D * P_D) / ((P_T_pos_given_not_D * P_not_D) + (P_T_pos_given_D * P_D))

    return P_D_given_T_pos
print(answer3())
```

Output:
```
0.16666666666666655
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

#### Let X be a discrete random variable that takes the values 1, 2, 3 with probabilities P(X=1)=0.2, P(X=2)=0.5, P(X=3)=0.3. 
-   Compute the expected value E[X].
-   Return your answer as a float.

```python
def ans4():

    # Computing the expected value E[X]
    E_X = ((1 * .2) + (2 * .5) + (3 * .3))

    # Return the answer as a float
    return float(E_X)

# You can test the function like this:
print(ans4())
    
```

Output:
```
2.0999999999999996
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

####  5. Let X be a random variable as in Question 4. Compute the variance Var(X).

-    Return your answer as a float.

```python
def ans5():
    values = [1, 2, 3]
    probabilities = [0.2, 0.5, 0.3]

    # Computing the expected value E[X]
    E_X = ((1 * .2) + (2 * .5) + (3 * .3))

    # Computing the variance Var(X)
    variance = sum(p * ((v - E_X) ** 2) for v, p in zip(values, probabilities))

    print(E_X)
    print(variance)

    return float(variance)  # Ensure it returns a float

```

Output:
```
2.0999999999999996
0.49
0.49
```

```python
ans5()
```

Output:
```
2.0999999999999996
0.49
0.49
```

```python
grader.check("q5")
```

Output:
```
q5 results:
    q5 - 1 result:
        ❌ Test case failed
        Trying:
            assert type(ans5() == float)
        Expecting nothing
        ok
        Trying:
            assert abs(ans5() - 0.49) < 1e-09
        Expecting nothing
        ok
        Trying:
            assert 0.49 < ans5() < 0.491
        Expecting nothing
        **********************************************************************
        Line 3, in q5 0
        Failed example:
            assert 0.49 < ans5() < 0.491
        Exception raised:
            Traceback (most recent call last):
              File "C:\Program Files\WindowsApps\PythonSoftwareFoundation.Python.3.12_3.12.1776.0_x64__qbz5n2kfra8p0\Lib\doctest.py", line 1368, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q5 0[2]>", line 1, in <module>
                assert 0.49 < ans5() < 0.491
                       ^^^^^^^^^^^^^^^^^^^^^
            AssertionError
```

#### 6. A fair coin is tossed 10 times. What is the probability of getting exactly 6 heads?
    Return your answer as a float.

```python
from math import comb
def ans6():
# n=10 is the number of trials (coin tosses),
# k=6 is the number of successful trials (heads),
# p=0.5 is the probability of getting heads in a single toss (since it's a fair coin).
    n = 10
    k =6
    p = 0.5

    binomial_coeff = comb(n,k)

    probability = binomial_coeff * (p ** k) * (1 - p) ** (n-k)
    return float(probability)
```

```python
ans6()
```

Output:
```
0.205078125
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

#### 7. A radioactive source emits particles following a Poisson distribution with an average rate of λ = 4 particles per minute. 
    What is the probability that exactly 6 particles are emitted in a given minute?

    Return your answer as a float.

```python
import math
def ans7():
    λ = 4
    k = 6
    probability = (λ** k * math.exp(-λ)) / math.factorial(k)
    return float(probability)
print(ans7())
```

Output:
```
0.1041956345670211
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

```python
pip install scipy

```

Output:
```
Defaulting to user installation because normal site-packages is not writeable
Collecting scipy
  Downloading scipy-1.14.1-cp312-cp312-win_amd64.whl.metadata (60 kB)
Requirement already satisfied: numpy<2.3,>=1.23.5 in c:\users\danst\appdata\local\packages\pythonsoftwarefoundation.python.3.12_qbz5n2kfra8p0\localcache\local-packages\python312\site-packages (from scipy) (2.1.1)
Downloading scipy-1.14.1-cp312-cp312-win_amd64.whl (44.5 MB)
   ---------------------------------------- 0.0/44.5 MB ? eta -:--:--
   -- ------------------------------------- 2.4/44.5 MB 12.2 MB/s eta 0:00:04
   ---- ----------------------------------- 5.5/44.5 MB 14.0 MB/s eta 0:00:03
   ------- -------------------------------- 8.4/44.5 MB 14.0 MB/s eta 0:00:03
   ---------- ----------------------------- 11.5/44.5 MB 14.4 MB/s eta 0:00:03
   ------------- -------------------------- 15.5/44.5 MB 15.2 MB/s eta 0:00:02
   ---------------- ----------------------- 18.6/44.5 MB 15.2 MB/s eta 0:00:02
   -----------------
```

#### 8. Assuming that the heights of adult males are normally distributed with a mean of 70 inches and a standard deviation of 3 inches, what is the probability that a randomly selected male is taller than 75 inches?

    Return your answer as a float.

```python
from scipy.stats import norm
def ans8():
    mu = 70
    sigma = 3
    x = 75
    
    # Answer is 1 - cumulative distribution function for x <= 75
    cdf = norm.cdf(x, mu, sigma)
    probability = 1 - cdf
    return float(probability)
print(ans8())
```

Output:
```
0.047790352272814696
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

#### 9. Write a function to calculate the Jaccard index (also known as Jaccard similarity coefficient) between two sets.

    Example:
    Input: set1 = {1, 2, 3, 4, 5}, set2 = {2, 4, 6, 8}
    Output: 0.2857
    Explanation:
    Intersection: {2, 4}
    Union: {1, 2, 3, 4, 5, 6, 8}
    Jaccard Index = 2 / 7 ≈ 0.2857

```python
def ans9(set1: set, set2: set) -> float:
    # Jaccard index = intersection / union
    intersection = set1.intersection(set2)
    union = set1.union(set2)
    # print(intersection)
    # print(union)

    jaccard_index = len(intersection) / len(union)

    return jaccard_index

# set1 = {1, 2, 3, 4, 5}
# set2 = {2, 3, 4}
# print(ans9(set1, set2))

```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

#### 10. Write a function that computes the norms between two lists.  (1-norm, 2-norm, infinity-norm, euclidean distance)
    Return a dictionary with keys `1_norm`, `2_norm`, `inf_norm` and `euclidean`.
    Note: Convert all the norms to float while returning.

```python
def ans10(arr1: np.array, arr2: np.array) -> float:
    

    arr1 = np.array(arr1)
    arr2 = np.array(arr2)
    diff = arr1 - arr2

    norms = {
        "1_norm": float(np.linalg.norm(diff, ord=1)),
        "2_norm": float(np.linalg.norm(diff, ord=2)),
        "inf_norm": float(np.linalg.norm(diff, ord=np.inf)),
        "euclidean": float(np.linalg.norm(diff))
    }

    return norms

arr1 = [1, 2, 3]
arr2 = [2, 3, 4]
print(ans10(arr1, arr2))
```

Output:
```
{'1_norm': 3.0, '2_norm': 1.7320508075688772, 'inf_norm': 1.0, 'euclidean': 1.7320508075688772}
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
C:\Users\danst\AppData\Local\Packages\PythonSoftwareFoundation.Python.3.12_qbz5n2kfra8p0\LocalCache\local-packages\Python312\site-packages\otter\check\notebook.py:470: UserWarning: Could not locate a PDF to include
  warnings.warn("Could not locate a PDF to include")
File ~\AppData\Local\Packages\PythonSoftwareFoundation.Python.3.12_qbz5n2kfra8p0\LocalCache\local-packages\Python312\site-packages\nbformat\__init__.py:169, in read(fp, as_version, capture_validation_error, **kwargs)
    141 """Read a notebook from a file as a NotebookNode of the given version.
    142 
    143 The string can contain a notebook of any version.
   (...)
    165     The notebook that was read.
    166 """
    168 try:
--> 169     buf = fp.read()
    170 except AttributeError:
    171     with open(fp, encoding="utf8") as f:  # noqa: PTH123

File C:\Program Files\WindowsApps\PythonSoftwareFoundation.Python.3.12_3.12.1776.0_x64__qbz5n2kfra8p0\Lib\encodings\cp1252.py:23, in IncrementalDecoder.decode(self, input
```

Tests for q10 (points: None)
```
>>> assert type(ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))) == dict
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['1_norm'] == 9.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['inf_norm'] == 3.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['2_norm'] > 5.0
```


Tests for q2 (points: None)
```
>>> assert len(roll_die(10)) == 10
>>> assert all((1 <= x <= 6 for x in roll_die(100)))
```


Tests for q3 (points: None)
```
>>> assert answer3() > 0.15 and answer3() < 0.2
>>> assert type(answer3()) == float
```


Tests for q4 (points: None)
```
>>> assert type(ans4()) == float
>>> assert 2 < ans4() < 2.1
```


Tests for q5 (points: None)
```
>>> assert type(ans5() == float)
>>> assert abs(ans5() - 0.49) < 1e-09
>>> assert 0.49 < ans5() < 0.491
```


Tests for q6 (points: None)
```
>>> assert type(ans6()) == float
>>> assert 0.2 < ans6() < 0.21
>>> assert abs(ans6()) - 0.2 < 0.05
```


Tests for q7 (points: None)
```
>>> assert type(ans7()) == float
>>> assert 0.1 < ans7() < 0.11
>>> assert abs(ans7()) - 0.1 < 0.0042
```


Tests for q8 (points: None)
```
>>> assert type(ans8()) == float
>>> assert 0.04 < abs(ans8()) < 0.05
```


Tests for q9 (points: None)
```
>>> assert ans9({1, 2, 3, 4, 5}, {2, 3, 4}) == 0.6
>>> assert 0.4 < ans9({1, 2, 3, 4, 5, 6, 7}, {2, 3, 4}) < 0.5
>>> assert 0.2 < ans9({1, 2, 3, 4, 5, 6, 7, 8, 9}, {2, 3}) < 0.23
```
